## **About the data** 
---
The data set consists of 116,658 observations and 10 columns. It contains data of fifth-grade students, including their Math final exam grade.


**The data set is available in the folder data**

```python
print('a')
```

Output:
```
a
```

```python
import emoji
```

<!-- grade_id: 0-0, points: 0, grade: true, solution: true -->
```python
# Your libraries here
# YOUR CODE HERE
raise NotImplementedError()
```

Output:
```
NotImplementedError                       Traceback (most recent call last)
Cell In[1], line 3
      1 # Your libraries here
      2 # YOUR CODE HERE
----> 3 raise NotImplementedError()

NotImplementedError:
```

<!-- grade_id: 0-1-1, points: 0, grade: true, solution: false -->
```python
# TEST CELL, DONT MODIFY
assert len(df) == 116658
```

Output:
```
NameError                                 Traceback (most recent call last)
Cell In[8], line 2
      1 # TEST CELL, DONT MODIFY
----> 2 assert len(df) == 116658

NameError: name 'df' is not defined
```

### 1.2
Justify the choice of each descriptive statistic. What do they say about the data? Can you identify some irregularities? 

<!-- grade_id: 1-2, points: 5, grade: true, solution: true -->
YOUR ANSWER HERE

### 1.3
In a single figure, choose an appropriate type of graph for each feature and plot each feature individually.

<!-- grade_id: 1-3, points: 10, grade: true, solution: true -->
```python
#### GRADED CELL ####
### 1.3
def plot_features(df):
    """
    Plots all features individually in the same figure
    
    Parameters
    ----------
    df : DataFrame
         Containing all data
         
    Hint
    ------
    To have multiple plots in a single figure see pyplot.figure

    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
sample = df.sample(n=1000)
plot_features(sample)
```