## **About the data** 
---
The data set consists of 116,658 observations and 10 columns. It contains data of fifth-grade students, including their Math final exam grade.


**The data set is available in the folder data**

<!-- grade_id: 0-0, points: 0, grade: true, solution: true -->
```python
# Your libraries here
# YOUR CODE HERE
raise NotImplementedError()
```

Output:
```
NotImplementedError                      Traceback (most recent call last)
Cell In[1], line 3
      1 # Your libraries here
      2 # YOUR CODE HERE
----> 3 raise NotImplementedError()

NotImplementedError:
```

<!-- grade_id: 0-1-1, points: 0, grade: true, solution: false -->
```python
# TEST CELL, DONT MODIFY
assert len(df) == 116658
```

### 1.2
Justify the choice of each descriptive statistic. What do they say about the data? Can you identify some irregularities? 

<!-- grade_id: 1-2, points: 5, grade: true, solution: true -->
YOUR ANSWER HERE

### 1.3
In a single figure, choose an appropriate type of graph for each feature and plot each feature individually.

<!-- grade_id: 1-3, points: 10, grade: true, solution: true -->
```python
#### GRADED CELL ####
### 1.3
def plot_features(df):
    """
    Plots all features individually in the same figure
    
    Parameters
    ----------
    df : DataFrame
         Containing all data
         
    Hint
    ------
    To have multiple plots in a single figure see pyplot.figure

    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
sample = df.sample(n=1000)
plot_features(sample)
```

### 1.4
Explain your observations. How are the features distributed (poisson, exponential, gaussian, etc)? Can you visually identify outliers? 

<!-- grade_id: 1-4, points: 5, grade: true, solution: true -->
YOUR ANSWER HERE

<a id="section2"></a>
## **2 Data Cleaning** 
---

Using your findings from the previous section, carefully continue to explore the data set and do the following:

1. Create a function to handle the missing values
2. Justify your decisions to treat the missing values
3. Create a function to handle the inconsistent data
4. Justify your decisions to treat the inconsistent data


### 2.1
Create a function to handle the missing values

<!-- grade_id: 2-1, points: 10, grade: true, solution: true -->
```python
#### GRADED CELL ####
### 2.1
def handle_missing_values(df):
    """
    Identifies and removes all missing values

    Parameters
    ----------
    df : DataFrame
      Containing missing values

    Returns
    -------
    df : DataFrame
      Without missing values

    Hint:
    -----
    Try to understand the pattern in the missing values    
    """
    # YOUR CODE HERE
    raise NotImplementedError()
    return df

```

<!-- grade_id: 2-1-1, points: 5, grade: true, solution: false -->
```python
# TEST CELL, DONT MODIFY
assert len(handle_missing_values(df).columns) == 9

```

<!-- grade_id: 2-1-2, points: 5, grade: true, solution: false -->
```python
# TEST CELL, DONT MODIFY

```

```python
df = handle_missing_values(df)
print("length of the dataframe:", len(df))
df.head()
```

```python
# take a look at the new dataframe stats and compare it with the original
get_feature_stats(df)
```

### 2.2 
Justify your decisions to treat the missing values. Are there missing values? If so, how are the missing values encoded? Why are there missing values? Is there a pattern in the values missing?


<!-- grade_id: 2-2, points: 5, grade: true, solution: true -->
YOUR ANSWER HERE

### 2.3 
Create a function to handle the inconsistent data

<!-- grade_id: 2-3, points: 0, grade: true, solution: true -->
```python
#### GRADED CELL ####
### 2.3
def handle_inconsistent_data(df):
    """
    Identifies features with inconsistent data types and transforms features
    to the correct data type (numerical, object). 

    Parameters
    ----------
    df : DataFrame
      Containing inconsistent data

    Returns
    -------
    df : DataFrame
       With consistent data. All columns must be either numerical or categorical

    Hint:
    -----
    Don't forget to convert the features into the correct data type 
    """
    # YOUR CODE HERE
    raise NotImplementedError()
    return df
```

<!-- grade_id: 2-3-1, points: 10, grade: true, solution: false -->
```python
# TEST CELL, DONT MODIFY

```

<!-- grade_id: 2-3-2, points: 10, grade: true, solution: false -->
```python
# TEST CELL, DONT MODIFY

```

```python
df = handle_inconsistent_data(df)
print(len(df))
print(df.head())
print(get_feature_stats(df))
```

### 2.4 
Justify your decisions to treat the inconsistent data. Were there columns with inconsistent data types? How did you identify them? 
 

<!-- grade_id: 2-4, points: 5, grade: true, solution: true -->
YOUR ANSWER HERE

<a id="section3"></a>
## **3 Visualization** 
---

After cleaning the data, we can try to understand or extract insights from it. To do so, in this last section, you will do the following:
1. Create a function to show the relationship between numerical features.
2. Interpret your findings. What is correlation useful for? What insights can you get from it? 
3. Select an appropriate type of graph to explore the relationship between grade, school group, and any other meaningful feature
4. Interpret your findings. What are some factors that seem to influence the grade of the students? Which features do not seem to affect the outcome?


### 3.1 
Create a function to show the linear correlation between features.

<!-- grade_id: 3-1, points: 5, grade: true, solution: true -->
```python
#### GRADED CELL ####
### 3.1
import seaborn as sns
def plot_correlation(df):
    """
    Builds upper triangular heatmap with pearson correlation between numerical variables

    Instructions
    ------------
    The plot must have:
    - An appropiate title
    - Only upper triangular elements
    - Annotated values of correlation coefficients rounded to three significant 
    figures
    - Negative correlation must be blue and possitive correlation red. 

    Parameters
    ----------
    df : DataFrame with data


    """
    # YOUR CODE HERE
    raise NotImplementedError()

```

```python
plot_correlation(df)
```

### 3.2
Interpret your findings. What is correlation useful for? What insights can you get from it? 


<!-- grade_id: 3-3, points: 5, grade: true, solution: true -->
YOUR ANSWER HERE

### 3.3
Select an appropriate type of graph to explore the relationship between grade, school group, and any other meaningful feature.


<!-- grade_id: 3-4, points: 5, grade: true, solution: true -->
```python
#### GRADED CELL ####
### 3.3
def plot_grades(df):
    """
    Visualizes the relationship between grade,  school group and other meaningful
    feature

    Parameters
    ----------
    df : DataFrame with data

    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
plot_grades(df)
```

### 3.4
Interpret your findings. What are some factors that seem to influence the grade of the students? Which features do not seem to affect the outcome?

<!-- grade_id: 3-5, points: 5, grade: true, solution: true -->
YOUR ANSWER HERE